// // FIXME: Make me pass! Diff budget: 25 lines.

// // I AM NOT DONE

// #[derive(Debug,PartialEq)]
// enum Duration {
//     MilliSeconds(u64),
//     Seconds(u32),
//     Minutes(u16),
// }

// impl PartialEq for Duration {
    
// }

// // What traits does `Duration` need to implement?

// #[test]
// fn traits() {
//     assert_eq!(Duration::Seconds(120), Duration::Minutes(2));
//     assert_eq!(Duration::Seconds(420), Duration::Minutes(7));
//     assert_eq!(Duration::MilliSeconds(420000), Duration::Minutes(7));
//     assert_eq!(Duration::MilliSeconds(43000), Duration::Seconds(43));
// }

pub fn main() {}